// primitive_types2.rs
//
// Complétez le reste de la ligne où le code est manquant ! Pas d'indices, 
// il n'y a pas d'astuces, il faut juste s'habituer à les taper :)
//
// Exécutez `rustlings hint primitive_types2` ou utilisez la sous-commande `hint`
// de `watch` pour obtenir une indication.

// I AM NOT DONE

fn main() {
    // Characters (`char`)

    // Notez les apostrophes, ils sont différents des guillemets
    // que vous avez pu voir.
    let my_first_initial = 'C';
    if my_first_initial.is_alphabetic() {
        println!("Alphabetical!");
    } else if my_first_initial.is_numeric() {
        println!("Numerical!");
    } else {
        println!("Neither alphabetic nor numeric!");
    }

    let // Terminez cette ligne comme dans l'exemple ! Quel est ton caractère préféré ?
    // Essayez une lettre, essayez un chiffre, essayez un caractère spécial, 
    // essayez un caractère d'une autre langue que la vôtre, essayez un emoji!
    if your_character.is_alphabetic() {
        println!("Alphabetical!");
    } else if your_character.is_numeric() {
        println!("Numerical!");
    } else {
        println!("Neither alphabetic nor numeric!");
    }
}
