// tests1.rs
//
// Les tests sont importants pour s'assurer que votre code fait ce que vous pensez
// qu'il devrait faire. Les tests peuvent être exécutés sur ce fichier avec la commande suivante :
// rustlings run tests1
//
// Ce test a un problème -- faites-le compiler ! Faites passer le test !
// Faites échouer le test !
//
// Exécutez `rustlings hint tests1` ou utilisez la sous-commande `hint`
// de `watch` pour obtenir une indication.

// I AM NOT DONE

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert() {
        assert!();
    }
}
