// strings1.rs
//
// Faites-moi compiler sans changer la signature de la fonction !
//
// Exécutez `rustlings hint strings1` ou utilisez la sous-commande `hint`
// de `watch` pour obtenir une indication.

// I AM NOT DONE

fn main() {
    let answer = current_favorite_color();
    println!("My current favorite color is {}", answer);
}

fn current_favorite_color() -> String {
    "blue"
}
