// enums1.rs
//
// Pas d'indices cette fois ! ;)
//
// Exécutez `rustlings hint enums1` ou utilisez la sous-commande `hint`
// de `watch` pour obtenir une indication.

// I AM NOT DONE

#[derive(Debug)]
enum Message {
    // TODO: définir quelques types de messages, comme indiqué ci-dessous
}

fn main() {
    println!("{:?}", Message::Quit);
    println!("{:?}", Message::Echo);
    println!("{:?}", Message::Move);
    println!("{:?}", Message::ChangeColor);
}
